# Assignment2
## k-Nearest-Neighbours algorithm

k-nearest neighbors (KNN) algorithm is a type of supervised ML algorithm which can be used for both classification as well as regression predictive problems. However, it is mainly used for classification predictive problems in industry.

**Deadline August 26, 2021**

###  By: "YOUR NAME and ID"

Here, the necessary  python packages and dataset is pre-loaded for you.

```python
import numpy as np
import pandas as pd
import seaborn as sns
from sklearn import datasets
import matplotlib.pyplot as plt
from sklearn.neighbors import KNeighborsClassifier
from sklearn.decomposition import PCA
sns.set_style('whitegrid')
df = sns.load_dataset('iris')
df.head()
```

By using a label encoding  preprocessing technique, encode the species column and make the dataset ready for knn algorithm 

<!-- grade_id: preprocessing, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: preprocessingA, points: 10, grade: true, solution: false -->
```python
assert df.species.sum() == 150
```

After preparing the feature vectors as "X" and the target variable as "y" for classification, fit a k-Nearest-Neighbours(knn) classifier with n_neighbors equal to 50.

<!-- grade_id: KNN, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

Xv = X.values.reshape(-1,1)
h = 0.02
x_min, x_max = Xv.min(), Xv.max() + 1
y_min, y_max = y.min(), y.max() + 1
xx, yy = np.meshgrid(np.arange(x_min, x_max, h),
                     np.arange(y_min, y_max, h))

z = knn.predict(np.c_[xx.ravel(), yy.ravel()])
z = z.reshape(xx.shape)
```

<!-- grade_id: cell-7f962c2900d6f8f1, points: 10, grade: true, solution: false -->
```python
assert z.argmax() == 298
```

```python


fig = plt.figure(figsize=(8,5))

ax = plt.contourf(xx, yy, z, cmap = 'afmhot', alpha=0.3, 
                  );
plt.scatter(X.values[:, 0], X.values[:, 1], c=y, s=40, 
            alpha=0.9, edgecolors='k');
```

## PCA

Principal component analysis is a fast and flexible unsupervised method for dimensionality reduction in data. You have to fit 2 component PCA model to your preloaded dataset with 4 features. Ensure that the correct attribute is printed. Hint: Look at the format of the output. 


```python
[3.66123805 0.03604607]
```

<!-- grade_id: pcamodel, points: 0, grade: false, solution: true -->
```python
model = PCA(n_components=2)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: pcaW, points: 5, grade: true, solution: false -->
```python
assert model.get_covariance().sum() == 6.288502908277401
```

```python

```