## <center>Assignment 3</center>

### Deadline Thu Sep 17 ,  2021



## Problem setup

#### Problem

In  assignment 3 a binary classification is to assign customers into good or bad categories based on  prediction. Thus, you will build models and answer questions using real word data to  predict whether the customer repay their credit within 90 days. 

#### Data description

| Feature | Variable Type | Value Type | Description |
|:--------|:--------------|:-----------|:------------|
| age | Input Feature | integer | Customer age |
| DebtRatio | Input Feature | real | Total monthly loan payments (loan, alimony, etc.) / Total monthly income percentage |
| NumberOfTime30-59DaysPastDueNotWorse | Input Feature | integer | The number of cases when client has overdue 30-59 days (not worse) on other loans during the last 2 years |
| NumberOfTimes90DaysLate | Input Feature | integer | Number of cases when customer had 90+dpd overdue on other credits |
| NumberOfTime60-89DaysPastDueNotWorse | Input Feature | integer | Number of cased when customer has 60-89dpd (not worse) during the last 2 years |
| NumberOfDependents | Input Feature | integer | The number of customer dependents |
| SeriousDlqin2yrs | Target Variable | binary: <br>0 or 1 | Customer hasn't paid the loan debt within 90 days |


```python
# Disable warnings in Anaconda
import warnings
warnings.filterwarnings('ignore')
import numpy as np
import pandas as pd
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import GridSearchCV, StratifiedKFold
from sklearn.ensemble import RandomForestClassifier
from sklearn.ensemble import BaggingClassifier
from sklearn.model_selection import cross_val_score, RandomizedSearchCV
from sklearn.preprocessing import StandardScaler
from sklearn.ensemble import RandomForestClassifier
sns.set()
```

```python
from matplotlib import rcParams
rcParams['figure.figsize'] = 11, 8
#the function that will replace NaN values with the median for each column.
def fill_nan(table):
    for col in table.columns:
        table[col] = table[col].fillna(table[col].median())
    return table 
```

Now, read the data:

```python
data = pd.read_csv('data/credit.csv')
data.head()
```

Look at the variable types:

```python
data.dtypes
```

Check the class balance:

```python
ax = data['SeriousDlqin2yrs'].hist(orientation='horizontal', color='red')
ax.set_xlabel("number_of_observations")
ax.set_ylabel("unique_value")
ax.set_title("Target distribution")

print('Distribution of the target:')
data['SeriousDlqin2yrs'].value_counts()/data.shape[0]
```

Separate the input variable names by excluding the target:

```python
independent_columns_names = [x for x in data if x != 'SeriousDlqin2yrs']
independent_columns_names
```

Apply the function to replace *NaN* values:

```python
table = fill_nan(data)
```

Separate the target variable and input features:

```python
X = table[independent_columns_names]
y = table['SeriousDlqin2yrs']
```

## Bootstrapping

**Question 1.** Make an interval estimate of the average age for the customers who delayed repayment at the 90% confidence level. Use the example from the article as reference, if needed. Also, use `np.random.seed(0)` as before. What is the resulting interval estimate?

1. 52.59 – 52.86
2. 45.71 – 46.13
3. 45.68 – 46.17
4. 52.56 – 52.88

<!-- grade_id: cell-c31b556c6dab3590, points: 0, grade: false, solution: true -->
```python

def get_bootstrap_samples(data, n_samples):
    """Generate samples using bootstrapping."""
    indices = np.random.randint(0, len(data), (n_samples, len(data)))
    samples = data[indices]
    return samples

def stat_intervals(stat, alpha):
    """Make an interval estimate."""
    boundaries = np.percentile(stat, [100 * alpha / 2., 100 * (1 - alpha / 2.)])
    return boundaries

# Save the ages of those who let a delay
churn = data[data['SeriousDlqin2yrs'] == 1]['age'].values

# Set the random seed for reproducibility 
np.random.seed(0)

# Generate bootstrap samples and calculate the mean for each sample
# YOUR CODE HERE
raise NotImplementedError()
# Print the interval estimate for the sample means
print("Mean interval", stat_intervals(churn_mean_scores, 0.1))

```

<!-- grade_id: cell-dd62792d8a6b8909, points: 5, grade: true, solution: false -->
```python
assert stat_intervals(churn_mean_scores, 0.1).sum() >= 90
```

## Logistic regression
Now, we will create a `LogisticRegression` model and use `class_weight='balanced'` to make up for our unbalanced classes.

```python
lr = LogisticRegression(random_state=5, class_weight='balanced')
```

Let's try to find the best regularization coefficient, which is the coefficient `C` for logistic regression. Then, we will have an optimal model that is not overfit and is a good predictor of the target variable.

```python
parameters = {'C': (0.0001, 0.001, 0.01, 0.1, 1, 10)}
```

In order to find the optimal value of `C`, let's apply stratified 5-fold validation and look at the *ROC AUC* against different values of the parameter `C`. Use the `StratifiedKFold` function for this: 

```python
skf = StratifiedKFold(n_splits=5, shuffle=True, random_state=5)
```

One of the important metrics of model quality is the *Area Under the Curve (AUC)*. *ROC AUC* varies from 0 to 1. The closer ROC AUC is to 1, the better the quality of the classification model. impliment a **Grid Search** with the scoring metric **"roc_auc"** for the parameter `C` and **StratifiedKFold** cross-validation.

<!-- grade_id: cell-c92fc4d555c1d496, points: 0, grade: false, solution: true -->
```python
## use variable name grid_search
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5f1a18ec19f948fb, points: 10, grade: true, solution: false -->
```python
assert grid_search.best_params_['C'] == 0.001
```

Can we consider the best model stable? The model is *stable* if the standard deviation on validation is less than 0.5%. Save the *ROC AUC* value of the best model; it will be useful for the following tasks.

1. Yes
2. No

```python
grid_search.cv_results_['std_test_score'][1]
```

```python
print(grid_search.best_score_*100)
```

## Feature importance with Logistic Regression

 *Feature importance* is defined by the absolute value of its corresponding coefficient. First, you need to normalize all of the feature values so that it will be valid to compare them and use logistic regression to fit the X,y dataset. What is the most important feature for the best logistic regression model?

1. age
2. NumberOfTime30-59DaysPastDueNotWorse
3. DebtRatio
4. NumberOfTimes90DaysLate
5. NumberOfTime60-89DaysPastDueNotWorse
6. MonthlyIncome
7. NumberOfDependents

<!-- grade_id: cell-f80c551c77f4c55e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
pd.DataFrame({'feat': independent_columns_names,
              'coef': lr.coef_.flatten().tolist()}).sort_values(by='coef', ascending=False)
```

<!-- grade_id: cell-9ca9256f157f5fd0, points: 10, grade: true, solution: false -->
```python
a =pd.DataFrame({'feat': independent_columns_names,
              'coef': lr.coef_.flatten().tolist()}).sort_values(by='coef', ascending=False)
assert a.coef.sum() == 0.9312707899227823
```

Calculate how much `DebtRatio` affects our prediction using the [softmax function](https://en.wikipedia.org/wiki/Softmax_function). What is its value?

1. 0.38
2. -0.02
3. 0.11
4. 0.24

```python
print((np.exp(lr.coef_[0]) / np.sum(np.exp(lr.coef_[0])))[2])
```

 Let's see how we can interpret the impact of our features. For this, recalculate the logistic regression with absolute values, that is without scaling. Next, modify the customer's age by adding 20 years, keeping the other features unchanged. How many times will the chance that the customer will not repay their debt increase?

1. -0.01
2. 0.70
3. 8.32
4. 0.66

```python
np.exp(lr.coef_[0][0]*20)
```

## Random Forest

Import the Random Forest classifier:

```python
from sklearn.ensemble import RandomForestClassifier
```

Initialize Random Forest with 100 trees and balance target classes:

```python
rf = RandomForestClassifier(n_estimators=100, n_jobs=-1, random_state=42, 
                            class_weight='balanced')
```

We will search for the best parameters among the following values:

```python
parameters = {'max_features': [1, 2, 4], 'min_samples_leaf': [3, 5, 7, 9], 'max_depth': [5,10,15]}
```

Also, we will use the stratified k-fold validation again. You should still have the `skf` variable.

How much higher is the **ROC AUC** of the best random forest model than that of the best logistic regression on validation?

1. 4%
2. 3%
3. 2%
4. 1%

<!-- grade_id: cell-9bacc0ebeb71dafc, points: 0, grade: false, solution: true -->
```python
%%time
# use rf_grid_search as a variable name of Grid Seard on random forest  model 
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-5994cf48838c9c04, points: 10, grade: true, solution: false -->
```python
assert rf_grid_search.best_score_ - grid_search.best_score_ == 0.026820670654852385
```

What feature has the weakest impact in the Random Forest model?

1. age
2. NumberOfTime30-59DaysPastDueNotWorse
3. DebtRatio
4. NumberOfTimes90DaysLate
5. NumberOfTime60-89DaysPastDueNotWorse
6. MonthlyIncome
7. NumberOfDependents

<!-- grade_id: cell-e6d4d3333515ec3c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

What is the most significant advantage of using *Logistic Regression* versus *Random Forest* for this problem?

1. Spent less time for model fitting;
2. Fewer variables to iterate;
3. Feature interpretability;
4. Linear properties of the algorithm.

## Bagging

Import modules and set up the parameters for bagging:

```python
from sklearn.ensemble import BaggingClassifier
from sklearn.model_selection import cross_val_score, RandomizedSearchCV

parameters = {'max_features': [2, 3, 4], 'max_samples': [0.5, 0.7, 0.9], 
              'base_estimator__C': [0.0001, 0.001, 0.01, 1, 10, 100]}
```

Fit a bagging classifier with `random_state=42`. For the base classifiers, use 100 logistic regressors and use `RandomizedSearchCV` instead of `GridSearchCV`. It will take a lot of time to iterate over all 54 variants, so set the maximum number of iterations for `RandomizedSearchCV` to 20. Don't forget to set the parameters `cv` and `random_state=1`. What is the best *ROC AUC* you achieve?

1. 80.75%
2. 80.12%
3. 79.62%
4. 76.50%

<!-- grade_id: cell-30b3e3f380ce3dc0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-267a979097fe6571, points: 10, grade: true, solution: false -->
```python
assert r_grid_search.best_score_ == 0.8087951623513827
```

```python

```